```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

# Lab 5: Data Cleaning and DataFrames


To receive credit for a lab, answer all questions correctly and submit before the deadline.

You must submit this assignment to Gradescope by the on-time deadline. We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline. This way, you will have ample time to contact staff for submission support. 

```python
import pandas as pd
import numpy as np
import zipfile
import matplotlib
import matplotlib.pyplot as plt
from pathlib import Path

plt.rcParams['figure.figsize'] = (8, 5)
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />


## Part 1: Tricky Pandas


<br>

### Question 1

Sometimes, `pandas` give you weird output that you may not expect.  The next series of question walks you through a few examples that might surprise you. 

<br>

---


#### Question 1a

For this question, you will define a function, `trick_1a` that returns a number that is the answer to a multiple-choice question.  You may need to experiment with DataFrames to arrive at your answers. 

The function `trick_1a` should not take any arguments. In the function you will: 

* Create a DataFrame named `trick1A` that has three columns labeled `Name`, `Name`, and `Age`. `trick1A` should have 3 rows of any values (this is up to you).
  
* Save the DataFrame `trick1A` to a `.csv` file called `trick1A.csv` without the index.

* Create another DataFrame, named `trick1B`, by reading in the file `trick1A.csv`.

Your function should return `1`, `2`, or `3` by observing the following. 

1. It was not possible to create a DataFrame with duplicate columns.
2. `trickA` and `trickB` have the same column names.
3. `trickA` and `trickB` have different column names.

Note, your function should only have three lines of code and use only the following parts of `pandas`: 
* [DataFrame constructor](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.html#pandas.DataFrame)
* [`to_csv` method](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.to_csv.html#pandas.DataFrame.to_csv)
* [`read_csv` function](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html#pandas.read_csv)  

```python
def trick_1a(): 
    # Construct DataFrame tricky1A
    tricky1A = pd.DataFrame(...)
    # Save the DataFrame tricky1A to a file tricky1A.csv 
    ...
    # Create DataFrame tricky1B by reading in file tricky1A.csv
    tricky1B = ...

    return None
```

```python
grader.check("q1a")
```

<br>

---

#### Question 1b

For this question, you will define a function, `trick_1b` that is to do the same task as the function description in Question 1a. 

Now, your function can still use the following parts of `pandas`: 

* [DataFrame constructor](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.html#pandas.DataFrame),
* [`to_csv` method](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.to_csv.html#pandas.DataFrame.to_csv), and
* [`read_csv` function](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html#pandas.read_csv)

But, it can also use other methods and functions after the constructor, e.g., 
* ['rename' method](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.rename.html#pandas.DataFrame.rename)

```python
def trick_1b(): 
    # Construct DataFrame tricky1A
    tricky1A = pd.DataFrame(...)
    # Save the DataFrame tricky1A to a file tricky1A.csv 
    ...
    # Create DataFrame tricky1B by reading in file tricky1A.csv
    tricky1B = ...

    return None
```

```python
grader.check("q1b")
```

<br> 

---

#### Question 1c 

Create a function `trick_1c` that does not take any arguments. 

Use the cell below to experiment with the following: 

* Create a DataFrame named `bools` that has four columns: `True`, `True`, `False`, `False`. Each column name should be Boolean.
* `bools` should have 4 rows; the values are up to you.
* Predict the shape of the DataFrame that results by running each of the three lines of code below. Pick a corresponding answer from the given list. Your function should return a list with three numbers, one for each line.
* You should be able to answer without running any code, but feel free to run code to check your answer.

```python
bools[True]
bools[[True, True, False, False]]
bools[[True, False]]
```

Answer choices:
1. DataFrame: 2 columns, 1 row
2. DataFrame: 2 columns, 2 rows
3. DataFrame: 2 columns, 3 rows
4. DataFrame: 2 columns, 4 rows
5. DataFrame: 3 columns, 1 rows
6. DataFrame: 3 columns, 2 rows
7. DataFrame: 3 columns, 3 rows
8. DataFrame: 3 columns, 4 rows
9. DataFrame: 4 columns, 1 rows
10. DataFrame: 4 columns, 2 rows
11. DataFrame: 4 columns, 3 rows
12. DataFrame: 4 columns, 4 rows
13. Error

**Hints:** It's recommended to refer to the pd.DataFrame documentation for information on the parameters `data` and `columns`.


```python
# Cell to experiment with 

...
```

```python
def trick_1c(): 
    # Return a list with three numbers, one for answer for each line.
    return None
trick_ans = trick_1c()
```

```python
grader.check("q1c")
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 2: Summary Statistics 

In this question you will define two general purpose functions that make it easy to qualitatively assess the contents of a DataFrame.


<br>

---

### Question 2

Complete the implementation of the function `population_stats`, which takes in a DataFrame `df` and returns a DataFrame indexed by the columns of `df`, with the following columns:
* `'num_nonnull'`, which contains the number of non-null entries in each column.
* `'prop_nonnull'`, which contains the proportion of entries in each column that are non-null.
* `'num_distinct'`, which contains the number of distinct non-null entries in each column.
* `'prop_distinct'`, which contains the proportion of non-null entries that are distinct in each column.
       
For example, if `df` has a column named `'ages'` with the following elements (note that `np.NaN` is a null value):
       
```py
[2, 2, 2, np.NaN, 5, 7, 5, 10, 11, np.NaN]
```

Then:
- `'num_nonnull'` is 8, and `'prop_nonnull'` is $\frac{8}{10}$ = 0.8.
- There are six distinct entries, `[2, 5, 7, 10, 11, np.NaN]`, but only 5 of them are non-null. So the number of distinct non-null entries, `'num_distinct'`, is 5.
- There are 5 distinct non-null entries, and there are 8 total non-null entries, so `'prop_distinct'` is $\frac{5}{8}$ = 0.625.

Putting it all together, `population_stats(df).loc['ages']` should be a Series containing the numbers 8, 0.8, 5, and 0.625.

<div class="alert alert-block alert-warning">
You should not use a for-loop to answer this question. 
</div>

<!-- grade_id: q3a-answer, points: 0, grade: false, solution: true -->
```python
def population_stats(df): 
    # Returns a DataFrame indexed by the columns of `df`, with the following columns:
    #  'num_nonnull', the number of non-null entries in each column.
    #  'prop_nonnull', the proportion of entries in each column that are non-null.
    #  'num_distinct', the number of distinct, non-null entries in each column.
    #  'prop_distinct', the proportion of distinct, non-null entries in each column.

    ...
    return df
    
pop_data = np.random.choice(range(10), size=(100, 4))
df_pop = pd.DataFrame(pop_data, columns = 'A B C D'.split())
out_pop = population_stats(df_pop)
out_pop.head()
```

```python
# Don't change this cell 
df_pop2 = pd.DataFrame({'A': [2, 2, 2, np.NaN, 5, 7, 5, 10, 11, np.NaN]})
out_pop2 = population_stats(df_pop2)
out_pop2.head()
```

```python
grader.check("q2")
```

<br>

---

### Question 3

Complete the implementation of the function `most_common`, which takes in a DataFrame `df` and a number `N` and returns a DataFrame of the `N` most-common values and their counts for each column of `df`. Any column with fewer than `N` distinct values should contain `np.NaN` in those entries.

For example, consider the DataFrame shown on the left. This DataFrame is a subset of `salaries`, a larger DataFrame containing information on employees in the City of San Diego. The subset below contains two of the original columns: `'Job Title'` which contains job titles for employees, and `'status'` which denotes whether the employee works a full time position (`'FT'`) or a part time position (`'PT'`). On the right, the return value of `most_common(salaries, N=5)` is shown.

You can assume that there are no ties in our hidden tests.

<table><tr>
    <td><img src="data/imgs/dataframe.png" width="80%"/></td>
    <td><img src="data/imgs/most_common.png" width="80%"/></td>
</tr></table>


***Hint***: You may find that initializing an empty DataFrame with `N` rows and adding columns to it is useful in your implementation.

<div class="alert alert-block alert-warning">

<b>Remember, the only question in this lab that you're allowed to use a loop in is Question 3 – that's this question.</b> Here, you may use a <code>for</code>-loop to loop over the columns in the input DataFrame (<code>df</code>), but not the rows. <b>If you use a <code>for</code>-loop or <code>while</code>-loop in any other question, you may lose points!</b>

</div>


```python
def most_common(df, N=5): 
    # Returns a DataFrame of the N most common values and their counts 
    #   for each column of the DataFrame df

    ...
    return None
    
```

```python
# don't change this cell 
common_data = np.random.choice(range(10), size=(100, 2))
common_df = pd.DataFrame(common_data, columns='A B'.split())
common_out = most_common(common_df, N=3)
common_out
```

```python
grader.check("q3")
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 3: Superheroes  

The questions below analyze a dataset of superheroes found in the `data` directory. One of the datasets lists the attributes of each superhero, while the other is a *Boolean* DataFrame describing which superheroes have which superpowers. Note, the datasets contain information on both **good** superheroes, as well as **bad** superheroes (AKA villains).


<br>

---

### Question 4


Let's start working with the `powers` dataset, which you can see in `data/superheroes_powers.csv`. 

Complete the implementation of the function `super_hero_powers`, which takes in a DataFrame like `powers` and returns a list with the following three entries:

1. The name of the superhero with the greatest number of superpowers.
2. The name of the second most common superpower among superheroes who can fly (the most common being `'Flight'` itself).
3. The name of the most common superpower among superheroes with only one superpower.

You should **not** be hard-coding your answers in this question; your function should work on any DataFrame similar to `powers`. In each case, you can assume the answer is unique.


```python
def super_hero_powers(df): 
    # Returns a list with the following entries: 
    #  - name of superhero with the greatest number of superpowers 
    #  - name of the second most common superpower among superheros that fly
    #  - name of most common superpower among superheroes with only one superpower

    ...
    return None
    

super_fp = Path('data') / 'superheroes_powers.csv'
powers = pd.read_csv(super_fp)
super_out = super_hero_powers(powers)
```

```python
grader.check("q4")
```

<br> 

--- 

### Question 5

You will use the dataset in `data/superheroes.csv` as a DataFrame. Your function `population_stats` from Question 2 will be called on the DataFrame. You should notice that there are very few actually null (`np.NaN`) values, but there are many entries that **should** be null, because they're missing.

Complete the implementation of the function `clean_heroes`, which takes in a DataFrame like the one created from `superheroes.csv` and returns a new DataFrame with all of the missing values replaced with `np.NaN`.

After cleaning the superheroes dataset with `clean_heroes`, you will run `population_stats` on it again. As a result of the cleaning, population_stats should show that there are more null values.

***Note***: Most of the work in this question is identifying how the missing values are stored in the DataFrame. The implementation of the function should only take one line.

<!-- grade_id: q4a-answer, points: 0, grade: false, solution: true -->
```python
def clean_heroes(df): 
    # Returns a DataFrame with all the missing values replaced with np.NaN 

    ...
    return None
    

superheroes_fp = Path('data') / 'superheroes.csv'
heroes = pd.read_csv(superheroes_fp, index_col=0)
print(population_stats(heroes))

clean_out = clean_heroes(heroes)
print(population_stats(clean_out))
```

```python
grader.check("q5")
```

<br>

---

### Question 6

Using the **cleaned** superhero data, `clean_out`, we will now generate some insights in the following questions. 

For each subquestion other than Question 6c and 6d, you should generate the answer using `pandas` methods and functions. 
For Question 6c and 6d, you can use your `pandas` knowledge to determine the information to answer the question, but you can hard-code in your final response. 



<br>

---

#### Question 6a

What is the name of the tallest `'Mutant'` with `'No Hair'`?


```python
tallest_no_hair = ...
    
tallest_no_hair
```

```python
grader.check("q6a")
```

<br>

---

#### Question 6b

Among the publishers who have more than 5 characters, which publisher has the highest proportion of human characters? If there is a tie, return the publisher whose name is first alphabetically. We define a character to be human if their `'Race'` is exactly the string `'Human'`; for instance, a `'Race'` of `'Human / Radiation'` is non-human for the purposes of this question.


```python

...
publisher_prop_human = ...
    
publisher_prop_human
```

```python
grader.check("q6b")
```

<br> 

--- 

#### Question 6c 

Among the characters whose `'Height'`s we know, who is taller on average – `'good'` characters or `'bad'` characters?

Use your `pandas` knowledge to determine the information to answer the question, but you can hard-code in your final response. 


```python
...
taller_chars = ...
taller_chars
```

```python
grader.check("q6c")
```

<br>

---

#### Question 6d 

Which publisher has a greater proportion of `'bad'` characters – `'Marvel Comics'` or `'DC Comics'`?

Use your `pandas` knowledge to determine the information to answer the question, but you can hard-code in your final response. 



```python
...
publish_more_bad = ...
publish_more_bad
```

```python
grader.check("q6d")
```

<br> 

--- 

#### Question 6e 

There is only one character that is **both** more than one standard deviation above the mean in height and more than one standard deviation below the mean in weight.  What is their name?

```python
...
tall_and_light = ...
tall_and_light
```

```python
grader.check("q6e")
```

<br> 

--- 

#### Question 6f 

Which `'Publisher'` that isn't `'Marvel Comics'` or `'DC Comics'` has the most characters? Consider all characters whose `'Publisher'` we know – that is, don't drop rows because they have null values in other columns.

```python
...
most_chars = ...
most_chars
```

```python
grader.check("q6f")
```

<br><br>

<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Congratulations! You have finished Lab 05!


Congrats! You are finished with this assignment.

**Important**: To make sure the test cases run correctly, click `Kernel>Restart & Run All` and make sure all of the test cases are still passing. Doing so will submit your code for you. 

If your test cases are no longer passing after restarting, it's likely because you're missing a variable, or the modifications that you'd previously made to your DataFrame are no longer taking place (perhaps because you deleted a cell). 

You may submit this assignment as many times as you'd like before the deadline.

**You must restart and run all cells before submitting. Otherwise, you may pass test cases locally, but not on our servers. We will not entertain regrade requests of the form, “my code passed all of my local test cases, but failed the autograder”.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: [1, 3])
```
>>> ans = trick_1a()
>>> ans == 1 or ans == 2 or ans == 3
True
```
```
>>> trick_1a() == 1
True
```


Tests for q1b (points: [1])
```
>>> ans = trick_1b()
>>> ans == 1 or ans == 2 or ans == 3
True
```


Tests for q1c (points: [1, 1])
```
>>> isinstance(trick_ans, list)
True
```
```
>>> isinstance(trick_ans[1], int)
True
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1, 4])
```
>>> out_pop.index.tolist() == ['A', 'B', 'C', 'D']
True
```
```
>>> out_pop.columns.tolist() == ['num_nonnull', 'prop_nonnull', 'num_distinct', 'prop_distinct']
True
```
```
>>> (out_pop['num_distinct'] <= 10).all()
True
```
```
>>> (out_pop['prop_nonnull'] == 1.0).all()
True
```
```
>>> out_pop2.index.tolist() == ['A']
True
```
```
>>> out_pop2.columns.tolist() == ['num_nonnull', 'prop_nonnull', 'num_distinct', 'prop_distinct']
True
```
```
>>> all(out_pop2.iloc[0] == [8.0, 0.8, 5.0, 0.625])
True
```


Tests for q3 (points: [1, 1, 1, 1])
```
>>> common_out.index.tolist() == [0, 1, 2]
True
```
```
>>> common_out.columns.tolist() == ['A_values', 'A_counts', 'B_values', 'B_counts']
True
```
```
>>> common_out['A_values'].isin(range(10)).all()
True
```
```
>>> common_out['B_values'].isin(range(10)).all()
True
```


Tests for q4 (points: [1, 1, 1])
```
>>> isinstance(super_out, list)
True
```
```
>>> len(super_out) == 3
True
```
```
>>> all([isinstance(x, str) for x in super_out])
True
```


Tests for q5 (points: [2, 2])
```
>>> sum(clean_out['Skin color'].isnull()) > sum(heroes['Skin color'].isnull())
True
```
```
>>> sum(clean_out['Weight'].isnull()) > sum(heroes['Weight'].isnull())
True
```


Tests for q6a (points: [1])
```
>>> isinstance(tallest_no_hair, str)
True
```


Tests for q6b (points: [1, 2])
```
>>> isinstance(publisher_prop_human, str)
True
```
```
>>> publisher_prop_human in ['Dark Horse Comics', 'DC Comics', 'George Lucas', 'Marvel Comics', 'Star Trek']
True
```


Tests for q6c (points: [1])
```
>>> taller_chars in ['good', 'bad']
True
```


Tests for q6d (points: [1, 1])
```
>>> isinstance(publish_more_bad, str)
True
```
```
>>> publish_more_bad in ['Marvel Comics', 'DC Comics']
True
```


Tests for q6e (points: [1])
```
>>> isinstance(tall_and_light, str)
True
```


Tests for q6f (points: [1, 1])
```
>>> isinstance(most_chars, str)
True
```
```
>>> most_chars not in ['Marvel Comics', 'DC Comics']
True
```
