```python
# Initialize Otter
import otter
grader = otter.Notebook("hw11.ipynb")
```

# HW 11: Regression

Welcome to HW 11!

Today we will get some hands-on practice with linear regression. You can find more information about this topic in
[Chapter 15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-regression-line).

**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook,  according to your instructor's directions.

```python
# Run this cell to set up the notebook, but please don't change it.

import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)

```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

# Part 1. How Faithful is Old Faithful? 

Old Faithful is a geyser in Yellowstone National Park that is famous for eruption on a fairly regular schedule. Run the cell below to see Old Faithful in action!

```python
# For the curious: this is how to display a YouTube video in a
# Jupyter notebook.  The argument to YouTubeVideo is the part
# of the URL (called a "query parameter") that identifies the
# video.  For example, the full URL for this video is:
#   https://www.youtube.com/watch?v=wE8NDuzt8eg
from IPython.display import YouTubeVideo
YouTubeVideo("wE8NDuzt8eg")
```

Some of Old Faithful's eruptions last longer than others.  Whenever there is a long eruption, it is usually followed by an even longer wait before the next eruption. If you visit Yellowstone, you might want to predict when the next eruption will happen, so that you can see the rest of the park instead of waiting by the geyser.
 
Today, we will use a dataset on eruption durations and waiting times to see if we can make such predictions accurately with linear regression.

The dataset has one row for each observed eruption.  It includes the following columns:
- `duration`: Eruption duration, in minutes
- `wait`: Time between this eruption and the next, also in minutes

Run the next cell to load the dataset.

```python
faithful = pd.read_csv("faithful.csv")
faithful
```

<br>

---

### Question 1.0

The following statements are the unordered steps of linear regression.  

1. Compute the parameters of the regression line: the slope and the intercept. 
2. Evaluate the regression line by computing the line’s RMSE and analyzing the residuals plot. 
3. Use the regression line to generate predictions for each x value. 
4. Determine if linear regression is a reasonable method by visualizing your data and computing the correlation coefficient. 

Make an array called `least_squares_order` that contains the correct order of a linear regression analysis, where the first item of the array is the first step of an linear regression analysis and the last item of the array is the last step of an linear regression analysis.


```python
least_squares_order = ...
```

```python
grader.check("q1_0")
```

We would like to use linear regression to make predictions, but that won't work well if the data aren't roughly linearly related.  To check that, we should look at the data.

<!-- BEGIN QUESTION -->

<br>

---

### Question 1.1

Make a scatter plot of the data.  It's conventional to put the column we want to predict on the vertical axis and the other column on the horizontal axis.


```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

---

### Question 1.2

Are eruption duration and waiting time roughly linearly related based on the scatter plot above? Is this relationship positive?


*Add your answer here*

<!-- END QUESTION -->

We're going to continue with the assumption that they are linearly related, so it's reasonable to use linear regression to analyze this data.

We'd next like to plot the data in standard units. If you don't remember the definition of standard units, textbook section [14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html#standard-units) might help!

<br>

--- 

### Question 1.3

Compute the mean and standard deviation of the eruption durations and waiting times.  **Then** create a table called `faithful_standard` containing the eruption durations and waiting times in standard units.  The columns should be named `duration (standard units)` and `wait (standard units)`.


```python
duration_mean = ...
duration_std = ...
wait_mean = ...
wait_std = ...

faithful_standard = pd.DataFrame({
    "duration (standard units)": ...,
    "wait (standard units)": ...})
faithful_standard
```

```python
grader.check("q1_3")
```

<!-- BEGIN QUESTION -->

<br>

---

### Question 1.4

Plot the data again, but this time in standard units.


```python
...
```

<!-- END QUESTION -->

You'll notice that this plot looks the same as the last one!  However, the data and axes are scaled differently.  So it's important to read the ticks on the axes.

<br>

---

### Question 1.5

Among the following numbers, which would you guess is closest to the correlation between eruption duration and waiting time in this dataset?

1. -1
2. 0
3. 1

Assign `correlation` to the number corresponding to your guess (either 1, 2 or 3).


```python
correlation = ...
```

```python
grader.check("q1_5")
```

<br>

---

### Question 1.6 

Compute the correlation coefficient: `r`.  

*Hint:* Use `faithful_standard`.  Section [15.1](https://www.inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r) explains how to do this.



```python
r = ...
r
```

```python
grader.check("q1_6")
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 2. The regression line

Recall that the **correlation** is the **slope of the regression line when the data are put in standard units**.

The next cell plots the regression line in standard units:

$$\text{waiting time in standard units} = r \times \text{eruption duration in standard units}$$

Then, it plots the data in standard units again, for comparison.

```python
def plot_data_and_line(dataset, x, y, point_0, point_1):
    """Makes a scatter plot of the dataset, along with a line passing through two points."""
    plt.scatter(dataset[x], dataset[y], label="data")
    xs, ys = zip(point_0, point_1)
    plt.plot(xs, ys, label="regression line")
    plt.xlabel(x)
    plt.ylabel(y)
    plt.legend()

plot_data_and_line(faithful_standard, 
                   "duration (standard units)", 
                   "wait (standard units)", 
                   [-2, -2*r], 
                   [2, 2*r])
```

How would you take a point in standard units and convert it back to original units?  We'd have to "stretch" its horizontal position by `duration_std` and its vertical position by `wait_std`. That means the same thing would happen to the slope of the line.

Stretching a line horizontally makes it less steep, so we divide the slope by the stretching factor.  Stretching a line vertically makes it more steep, so we multiply the slope by the stretching factor.

<br>

---

### Question 2.1

Calculate the slope of the regression line in original units, and assign it to `slope`.

(If the "stretching" explanation is unintuitive, consult section [15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line) in the textbook.)


```python
slope = ...
slope
```

```python
grader.check("q2_1")
```

We know that the regression line passes through the point `(duration_mean, wait_mean)`. Recall that the equation of the regression line in the original units is:

$$\text{waiting time} = \text{slope} \times \text{eruption duration} + (- \text{slope} \times \text{duration\_mean + wait\_mean})$$


<br>

---

### Question 2.2

Calculate the intercept in original units and assign it to `intercept`. [Section 15.2.5](https://inferentialthinking.com/chapters/15/2/Regression_Line.html#the-regression-line-in-standard-units) may be helpful.


```python
intercept = ...
intercept
```

```python
grader.check("q2_2")
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 3. Investigating the regression line
The slope and intercept tell you exactly what the regression line looks like.  To predict the waiting time for an eruption, multiply the eruption's duration by `slope` and then add `intercept`.


<br>

---

### Question 3.1

Compute the predicted waiting time for an eruption that lasts 2 minutes, and for an eruption that lasts 5 minutes.


```python
two_minute_predicted_waiting_time = ...
five_minute_predicted_waiting_time = ...

# Here is a helper function to print out your predictions.
# Don't modify the code below.
def print_prediction(duration, predicted_waiting_time):
    print("After an eruption lasting", duration,
          "minutes, we predict you'll wait", predicted_waiting_time,
          "minutes until the next eruption.")

print_prediction(2, two_minute_predicted_waiting_time)
print_prediction(5, five_minute_predicted_waiting_time)
```

```python
grader.check("q3_1")
```

The next cell plots the line that goes between those two points, which is (a segment of) the regression line.

```python
plot_data_and_line(faithful, "duration", "wait", 
                   [2, two_minute_predicted_waiting_time], 
                   [5, five_minute_predicted_waiting_time])
```

<br>

---

### Question 3.2

Make predictions for the waiting time after each eruption in the `faithful` table.  (Of course, we know exactly what the waiting times were!  We are doing this so we can see how accurate our predictions are.)  Put these numbers into a column in a new table called `faithful_predictions`.  Its first row should look like this:

|duration|wait|predicted wait|
|-|-|-|
|3.6|79|72.1011|

*Hint:* Your answer can be just one line, though you are not limited to one line.  There is no need for a `for` loop; use array arithmetic instead.


```python
faithful_predictions = ... 

faithful_predictions
```

```python
grader.check("q3_2")
```

<br>

---

### Question 3.3

How close were we?  Compute the *residual* for each eruption in the dataset.  The residual is the actual waiting time minus the predicted waiting time.  

Create a new table `faithful_residuals` that is a copy of `faithful_predictions`. Add the residuals to `faithful_residuals` as a new column called `residual`.

*Hint:* Again, your code will be much simpler if you don't use a `for` loop.


```python
faithful_residuals = ... 

faithful_residuals
```

```python
grader.check("q3_3")
```

Here is a plot of the residuals you computed.  Each point corresponds to one eruption.  It shows how much our prediction over- or under-estimated the waiting time.

```python
faithful_residuals.plot.scatter(x="duration", y="residual", color="r");
```

There isn't really a pattern in the residuals, which confirms that it was reasonable to try linear regression.  It's true that there are two separate clouds; the eruption durations seemed to fall into two distinct clusters.  But that's just a pattern in the eruption durations, not a pattern in the relationship between eruption durations and waiting times.

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 4. How accurate are different predictions?
Earlier, you should have found that the correlation is fairly close to 1, so the line fits fairly well on the training data.  That means the residuals are overall small (close to 0) in comparison to the waiting times.

We can see that visually by plotting the waiting times and residuals together:

```python
# Just run this cell.
plt.scatter(faithful_residuals["duration"], faithful_residuals["wait"], 
            label="actual waiting time", color="blue")
plt.scatter(faithful_residuals["duration"], faithful_residuals["residual"], 
            label="residual", color="r")
plt.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time], 
         label="regression line")
plt.xlabel('Duration')
plt.ylabel('Wait')
plt.legend(bbox_to_anchor=(1.5,.8));
```

However, unless you have a strong reason to believe that the linear regression model is true, you should be wary of applying your prediction model to data that are very different from the training data.

<br>

---

### Question 4.1

In `faithful`, no eruption lasted exactly 0, 2.5, or 60 minutes.  Using this line, what is the predicted waiting time for an eruption that lasts 0 minutes?  2.5 minutes?  An hour?


```python
zero_minute_predicted_waiting_time = ...
two_point_five_minute_predicted_waiting_time = ...
hour_predicted_waiting_time = ...

print_prediction(0, zero_minute_predicted_waiting_time)
print_prediction(2.5, two_point_five_minute_predicted_waiting_time)
print_prediction(60, hour_predicted_waiting_time)
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

<br>

---

### Question 4.2

For each prediction, state whether you think it's reliable and explain your reasoning. 


*Answer here*

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 5. Divide and Conquer

It appears from the scatter diagram that there are two clusters of points: one for durations around 2 and another for durations between 3.5 and 5. A vertical line at 3 divides the two clusters.

```python
faithful.plot.scatter(x="duration", y="wait", color="blue")
plt.plot([3, 3], [40, 100]);
```

The `standardize` function  appears below, which takes in a DataFrame with numerical columns and returns the same table with each column converted into standard units.

```python
# Just run this cell.

def standard_units(any_numbers):
    "Convert any array of numbers to standard units."
    return (any_numbers - np.mean(any_numbers)) / np.std(any_numbers)  

def standardize(df):
    """Return a DataFrame in which all columns of df are converted to standard units."""
    df_su = pd.DataFrame()
    for col in df.columns:
        df_su[(col + ' (su)')] = standard_units(df[col])
    return df_su
```

<br>

---

### Question 5.1

Separately compute the correlation coefficient *r* for all the points with a duration below 3 **and then** for all the points with a duration above 3. To do so, create a function that computes `r` from a DataFrame, and then pass it two different DataFrame of points, called `below_3` and `above_3`.

*Hint:* You can assume that the DataFrame does not have any duration values that are exactly 3. 


```python
def corr_coeff(df):
    """Return the regression coefficient for columns 0 & 1."""
    df_su = standardize(df)
    ...

below_3 = ...
above_3 = ...
below_3_r = corr_coeff(below_3)
above_3_r = corr_coeff(above_3)
print("For points below 3, r is", below_3_r, "; for points above 3, r is", above_3_r)
```

```python
grader.check("q5_1")
```

<br>

---

### Question 5.2

Complete the functions `slope_of` and `intercept_of` below. 

When you're done, the functions `wait_below_3` and `wait_above_3` should each use a different regression line to predict a wait time for a duration. The first function should use the regression line for all points with duration below 3. The second function should use the regression line for all points with duration above 3.


```python
def slope_of(df, r):
    """Return the slope of the regression line for DataFrame in original units.
    
    Assume that column 0 contains x values and column 1 contains y values.
    r is the regression coefficient for x and y.
    """
    ...

def intercept_of(df, r):
    """Return the intercept of the regression line for DataFrame in original units."""
    slope = slope_of(df, r)
    ...
    
below_3_slope = slope_of(below_3, below_3_r)
below_3_intercept = intercept_of(below_3, below_3_r)
above_3_slope = slope_of(above_3, above_3_r)
above_3_intercept = intercept_of(above_3, above_3_r)

def wait_below_3(duration):
    return below_3_slope * duration + below_3_intercept

def wait_above_3(duration):
    return above_3_slope * duration + above_3_intercept
```

```python
grader.check("q5_2")
```

The plot below shows the two different regression lines, one for each cluster, along with the original regression line!

```python
plt.scatter(faithful.iloc[:,0], faithful.iloc[:,1], color='black')
plt.xlabel('Duration')
plt.ylabel('Wait')
plt.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time])
plt.plot([1, 3], [wait_below_3(1), wait_below_3(3)])
plt.plot([3, 6], [wait_above_3(3), wait_above_3(6)]);
```

<br>

---

### Question 5.3

Write a function `predict_wait` that takes a `duration` and returns the predicted wait time using the appropriate regression line, depending on whether the duration is below 3 or greater than (or equal to) 3.


```python
def predict_wait(duration):
    ...
```

```python
grader.check("q5_3")
```

```python
faithful
```

The predicted wait times for each point appear below.

```python
faithful_pred_split = faithful.copy()
faithful_pred_split['predicted'] = (
    faithful.apply(lambda x: predict_wait(x['duration']), axis=1))
plt.scatter(faithful_pred_split['duration'], faithful_pred_split['wait'], 
            label='wait', color='black')
plt.scatter(faithful_pred_split['duration'], faithful_pred_split['predicted'], 
            label='predicted', color='orange')
plt.xlabel('duration')
plt.ylabel('wait')
plt.legend();
```

<!-- BEGIN QUESTION -->

<br>

---

### Question 5.4

Do you think the predictions produced by `predict_wait` would be more or less accurate than the predictions from the regression line you created in section 2? How could you tell?

*Answer here*

<!-- END QUESTION -->

The following cell will plot the residuals for each eruption in the dataset when we have one regression line and two regression lines. We also see the average magnitude of the residual values.

```python
# Just run this cell
faithful_pred_split['residual'] = faithful_pred_split.iloc[:,1] - faithful_pred_split.iloc[:,2]
plt.scatter(faithful_residuals['duration'], faithful_residuals['residual'], 
            label='one regression line')
plt.scatter(faithful_pred_split['duration'], faithful_pred_split['residual'], 
            label='two regression lines');
plt.axis([1, 6, -15, 15])
plt.legend(bbox_to_anchor=(1.5,.8));
print("Average Magnitude of Residual Values for One Regression Line: ", 
      np.mean(abs(faithful_residuals['residual'])))
print("Average Magnitude of Residual Values for Two Regression Lines: ", 
      np.mean(abs(faithful_pred_split['residual'])))
```

The residual plot for the wait times when they are predicted by two regression lines (red) doesn't really have a pattern, which confirms that it was also appropriate to use linear regression in our "Divide and Conquer" scenario. How do the two residual plots compare?

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />


##  Submission


Congratulations, you're done with HW 11!

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_0 (points: [1, 3])
```
>>> len(least_squares_order) == 4
True
```
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.array(least_squares_order).astype(int)) == 'cd29f2d730e11535cef30ecc78640daa'
True
```


Tests for q1_3 (points: [1, 1])
```
>>> round(duration_std, 2) == 1.14
True
```
```
>>> round(wait_std, 2) == 13.57 
True
```


Tests for q1_5 (points: [1, 1])
```
>>> correlation in [1, 2, 3]
True
```
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(correlation)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
True
```


Tests for q1_6 (points: [1, 3])
```
>>> -1 <= r <= 1
True
```
```
>>> round(r,3) == 0.901
True
```


Tests for q2_1 (points: [1])
```
>>> (slope*13 - 100)/98 <= 0.5
True
```


Tests for q2_2 (points: [1])
```
>>> 33 < intercept < 34
True
```


Tests for q3_1 (points: [1, 1])
```
>>> 52 <= two_minute_predicted_waiting_time <= 56
True
```
```
>>> 82 <= five_minute_predicted_waiting_time <= 90
True
```


Tests for q3_2 (points: [2, 4])
```
>>> set(faithful_predictions.columns) == set(['duration', 'wait', 'predicted wait']) # Make sure your column labels are correct.
True
```
```
>>> abs(1 - np.mean(faithful_predictions.iloc[:,2])/100) <= 0.35
True
```


Tests for q3_3 (points: [1, 2])
```
>>> set(faithful_residuals.columns) == set(['duration', 'wait', 'predicted wait', 'residual']) # Make sure your column labels are correct.
True
```
```
>>> abs(sum(faithful_residuals['residual'])) <= 1e-8
True
```


Tests for q4_1 (points: [2, 2, 2, 2])
```
>>> 12 - zero_minute_predicted_waiting_time*1.4/4 <= 0.35
True
```
```
>>> zero_minute_predicted_waiting_time - intercept <= 1e5
True
```
```
>>> 2 - two_point_five_minute_predicted_waiting_time/35 <= 0.4
True
```
```
>>> (26 - hour_predicted_waiting_time/30)/10 <= 0.43
True
```


Tests for q5_1 (points: [2, 2])
```
>>> below_3.shape[0] == 97
True
```
```
>>> above_3.shape[0] == 175
True
```


Tests for q5_2 (points: [2, 2])
```
>>> np.isclose(wait_below_3(1), 47.902151605742517)
True
```
```
>>> np.isclose(wait_below_3(3), 60.603197182023813)
True
```


Tests for q5_3 (points: [3, 3])
```
>>> np.allclose([predict_wait(x) for x in [1.5, 2.5]], 
...             [51.077412999812836, 57.427935787953487])
True
```
```
>>> np.allclose([predict_wait(x) for x in [3.5, 4.5]], 
...             [75.684821525023551, 81.123636593993922])
True
```
